# Example Assignment: Introduction to Python

**Course:** Introduction to Programming  
**Due Date:** March 15, 2026

## Instructions

Complete the following exercises. Write your code in the cells marked `YOUR CODE HERE` and do not modify the test cells.

**Grading:** This assignment is worth 15 points total.

---

## Problem 1: Basic Arithmetic (5 points)

Write a function `multiply(a, b)` that takes two numbers and returns their product.

**Example:**
```python
multiply(3, 4)  # Returns 12
multiply(-2, 5) # Returns -10
```

<!-- grade_id: multiply, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: multiply_tests, points: 5, grade: true, solution: false -->
```python
# Test cases for multiply function
assert multiply(3, 4) == 12, "Failed: 3 * 4 should equal 12"
assert multiply(-2, 5) == -10, "Failed: -2 * 5 should equal -10"
assert multiply(0, 100) == 0, "Failed: 0 * 100 should equal 0"
assert multiply(7, 1) == 7, "Failed: 7 * 1 should equal 7"
print("✓ All tests passed for Problem 1!")
```

---

## Problem 2: String Manipulation (5 points)

Write a function `reverse_string(s)` that takes a string and returns it reversed.

**Example:**
```python
reverse_string("hello")  # Returns "olleh"
reverse_string("Python") # Returns "nohtyP"
```

<!-- grade_id: reverse_string, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: reverse_string_tests, points: 5, grade: true, solution: false -->
```python
# Test cases for reverse_string function
assert reverse_string("hello") == "olleh", "Failed: reverse of 'hello' should be 'olleh'"
assert reverse_string("Python") == "nohtyP", "Failed: reverse of 'Python' should be 'nohtyP'"
assert reverse_string("") == "", "Failed: reverse of empty string should be empty string"
assert reverse_string("a") == "a", "Failed: reverse of single character should be same character"
print("✓ All tests passed for Problem 2!")
```

---

## Problem 3: List Operations (5 points)

Write a function `sum_even_numbers(numbers)` that takes a list of integers and returns the sum of only the even numbers.

**Example:**
```python
sum_even_numbers([1, 2, 3, 4, 5, 6])  # Returns 12 (2+4+6)
sum_even_numbers([1, 3, 5])           # Returns 0 (no even numbers)
sum_even_numbers([10, 20, 30])        # Returns 60
```

<!-- grade_id: sum_even, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: sum_even_tests, points: 5, grade: true, solution: false -->
```python
# Test cases for sum_even_numbers function
assert sum_even_numbers([1, 2, 3, 4, 5, 6]) == 12, "Failed: sum of evens in [1,2,3,4,5,6] should be 12"
assert sum_even_numbers([1, 3, 5]) == 0, "Failed: sum of evens in [1,3,5] should be 0"
assert sum_even_numbers([10, 20, 30]) == 60, "Failed: sum of evens in [10,20,30] should be 60"
assert sum_even_numbers([]) == 0, "Failed: sum of evens in empty list should be 0"
assert sum_even_numbers([2]) == 2, "Failed: sum of evens in [2] should be 2"
print("✓ All tests passed for Problem 3!")
```

---

## Submission

Make sure all tests pass before submitting. Your notebook is automatically saved to your Google Drive and shared with the instructor.

**Do not modify test cells** or you will receive zero credit for this assignment.